---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ejer 1

```yaml
type: NormalExercise
key: ce5fa0dab7
lang: r
xp: 100
skills: 1
```

Contexto del ejer

`@instructions`
- Instruction 1
- Instruction 2

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
4+
```

`@sample_code`
```{r}
4+x
```

`@solution`
```{r}
4+3
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Muy bien, sigue adelante!.")
```
